# Problem 0: Feedback on the course so far (*1 point*)

To start the exercise this week, we would like to ask for your feedback on the course thus far. Please provide anonymous feedback at the link below.

Course feedback link: [https://elomake.helsinki.fi/lomakkeet/102085/lomake.html](https://elomake.helsinki.fi/lomakkeet/102085/lomake.html)

On the feedback form, you will find a secret word. Enter that word as the character string variable `secret_word` in the Python cell below for one point in this week's exercise :).

**Your score on this problem will be based on following criteria**:

- Providing some course feedback using the link above
- Entering the secret word in the Python cell below
- Uploading your notebook to your GitHub repository for this week's exercise

<!-- grade_id: problem_1_ag_p0_feedback, points: 0, grade: false, solution: true -->
```python
# Enter the secret word below
secret_word = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_1_ag_p0_feedback_test1, points: 1, grade: true, solution: false -->
```python
# This test should work
print('The secret word is {0}!'.format(secret_word))

```

# Problem 2 - Ice flow in an open channel (*5 points*)

The goal of this problem is to calculate horizontal velocity profiles across a glacier for Newtonian and non-Newtonian fluid flow resulting from a pressure gradient.

**Notice**: Closely follow the instructions! For example, you should be sure to use **exactly** the same variable names mentioned in the instructions because your answers will be automatically graded, and the tests that grade your answers rely on following the same formatting or variable naming as in the instructions.

**Your score on this problem will be based on following criteria**:

- Creating a Python function for calculating the pressure gradient
- Creating a Python function for the channel flow velocity
- Calculating and plotting ice velocities in a channel
- Including comments that explain what most lines in the code do
- Uploading your notebook to your GitHub repository for this week's exercise

## Part 1: Background and getting started (*0 points*)

Before you start you should read through the [background and theory](https://introqg.github.io/site/lessons/L5/exercise-5-theory.html#problem-1) for this problem. The other parts of this problem will refer to items described in this theory text. You should also run the two Python cells below to configuring plotting for this exercise and load the modules we will be using.

- Read through the [background and theory](https://introqg.github.io/site/lessons/L5/exercise-5-theory.html#problem-1) for this problem
- Run the two Python cells below

```python
# Enable inline plotting
%matplotlib inline
```

```python
# Import NumPy and pandas
import numpy as np
import pandas as pd
```

## Part 2: A pressure gradient function?!? (*1 point*)

As mentioned in the description of the problem, we will be calculating velocities of ice at the surface as we cross a glacial valley. We can calculate the ice velocity using Equation 10 from the [background and theory](https://introqg.github.io/site/lessons/L5/exercise-5-theory.html#problem-1), which is also shown below

\begin{equation}
  \large
  u(y) = \frac{a}{n + 1} \left( \frac{p_{1} - p_{0}}{L} \right)^{n} \left[ \left( \frac{h}{2} \right)^{n + 1} - y^{n + 1} \right]
\end{equation}

where $a$ is a function of the viscosity and temperature (we will ignore temperature for now), $n$ is the power-law exponent, $\frac{p_{1} - p_{0}}{L}$ is the pressure gradient in the channel, $h$ is the channel width, and $y$ is array for distance from the center of the channel. The only problem here is that the value for the pressure gradient $\left(\frac{p_{1} - p_{0}}{L}\right)^{n}$ is not known. Thus, our first step is to create a function to estimate its value from some of the observed velocities for the Saskatchewan glacier.

- Create a function called `channel_pressure` that calculates the value of $\left(\frac{p_{1} - p_{0}}{L}\right)^{n}$ for a given value of the power-law exponent `n`.
    - You can solve for $\left(\frac{p_{1} - p_{0}}{L}\right)^{n}$ by rearranging Equation 10 (see above) to solve for $\left(\frac{p_{1} - p_{0}}{L}\right)^{n}$ and using the condition that the maximum flow velocity $u = 2.41 \times 10^{-6}$ m/s occurs at $y = 0$ (see Table 1 in Part 4).
    - Assume $a = 5 \times 10^{-24}$ Pa<sup>-3</sup> s<sup>-1</sup> and $h = 1400$ m.
    - Your function parameters should include
        - The power-law exponent `n`
        - The channel width `h`
        - The viscosity/temperature constant `a`
        - The channel velocity at y = 0 `u_max`

<!-- grade_id: problem_2_ag_p2_pressure, points: 0, grade: false, solution: true -->
```python
def channel_pressure(n, h=1400.0, a=5.0e-24, u_max=2.41e-6):
    """Calculate ice channel pressure gradient.
    
    Keyword arguments:
    n -- viscous flow power-law expoenent
    h -- channel width (units: m; default: 1400.0)
    a -- viscosity/temperature constant (units: Pa^-3 s^-1; default: 5.0e-24)
    u_max -- Approximate channel velocity at y = 0 (units: m/s; default: 2.41e-6)
    """
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_2_ag_p2_pressure_test1, points: 0.5, grade: true, solution: false -->
```python
# The visible tests below should work
from nose.tools import assert_equal

# Function test calculations
pressure1 = channel_pressure(n=3, h=1400.0, a=5.0e-24, u_max=2.41e-6)

# Print results
print('The pressure gradient for n=3 should be 8029987.50521. My pressure gradient value is {0:.5f}'.format(pressure1))

# Check the other test calculations
assert_equal(round(pressure1, 5), 8029987.50521)
```

<!-- grade_id: problem_2_ag_p2_pressure_test2, points: 0.5, grade: true, solution: false -->
```python
# The visible tests below should work

# Function test calculations
pressure2 = channel_pressure(n=6, h=1400.0, a=5.0e-24, u_max=2.41e-6)
pressure3 = channel_pressure(n=1.5, h=1400.0, a=5.0e-24, u_max=2.41e-6)

# Print results
print('The pressure gradient for n=6 should be 0.04097. My pressure gradient value is {0:.5f}'.format(pressure2))

# Check the other test calculations
assert_equal(round(pressure2, 5), 0.04097)
assert_equal(round(pressure3, 5), 92948406117.57527)
```

## Part 3: An ice flow velocity function (*1 point*)

Your next task is to create a function to calculate the velocity (Equation 10) of ice flowing in a channel.

- Create a function called `channel_velocity` that calculates the velocity of ice flowing in a channel using Equation 10 from the [background and theory](https://introqg.github.io/site/lessons/L5/exercise-5-theory.html#problem-1), which was shown in the previous part.
    - Your function should use the `channel_pressure` function created in the previous part.
    - As in the previous part you should assume $a = 5 \times 10^{-24}$ Pa<sup>-3</sup> s<sup>-1</sup>, $h = 1400$ m, and that the flow is symmetric about $y = 0$.
    - Your function parameters should include
        - The power-law exponent `n`
        - The channel width array `y`
        - The channel width `h`
        - The viscosity/temperature constant `a`
        - The channel velocity at y = 0 `u_max`

<!-- grade_id: problem_2_ag_p3_velocity, points: 0, grade: false, solution: true -->
```python
def channel_velocity(n, y, h=1400.0, a=5.0e-24, u_max=2.41e-6):
    """Calculate planform velocity of ice flowing in a channel.
    
    Keyword arguments:
    n -- viscous flow power-law expoenent
    y -- array of distances from the channel center (units: m)
    h -- arbitrary channel width (units: m; default: 1400.0)
    a -- viscosity/temperature constant (units: Pa^-3 s^-1; default: 5.0e-24)
    u_max -- Approximate channel velocity at y = 0 (units: m/s; default: 2.41e-6)
    """
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_2_ag_p3_velocity_test1, points: 0.5, grade: true, solution: false -->
```python
# The visible tests below should work
from nose.tools import assert_equal

# Fake data
y1 = np.array([-500.0, -250.0, 0.0, 200.0, 400.0, 600.0])
velocity1 = channel_velocity(2, y1)

# Print results
print('The channel velocity at y1[1] for n = 2 should be 2.30022e-06 m/s. My channel velocity is {0:.5e} m/s.'.format(velocity1[1]))

# Check the other test calculations
assert_equal(round(velocity1[-1], 8), 8.9e-07)
assert_equal(round(velocity1[3], 8), 2.35e-06)
```

<!-- grade_id: problem_2_ag_p3_velocity_test2, points: 0.5, grade: true, solution: false -->
```python
# The visible tests below should work

# Fake data
y2 = np.array([-300, -150.0, 0.0, 100.0, 200.0, 300.0])
velocity2 = channel_velocity(4, y2)

# Print results
print('The channel velocity at y2[1] for n = 4 should be 2.40891e-06 m/s. My channel velocity is {0:.5e} m/s.'.format(velocity2[1]))

# Check the other test calculations
assert_equal(round(velocity2[-1], 8), 2.38e-06)
assert_equal(round(velocity2[3], 8), 2.41e-06)
```

## Part 4: Calculating ice flow velocities (*1 point*)

Your next task is to calculate the flow velocity for ice flowing in a channel using your `channel_velocity` function.

- Create a new pandas DataFrame called `channel` to store channel velocities for 1400-m-wide ice channel
    - Use `y` as the index for your `channel` DataFrame
    - **Hint**: You should define $h = 1400$ m and assume that the flow is symmetric about $y = 0$ and define `y` with a length of 201
- Calculate the velocity of ice in a channel as a non-Newtonian fluid with power-law exponents of $n$ = 2, 3, 4, and 5.
    - As above, you can assume $a = 5 \times 10^{-24}$ Pa<sup>-3</sup> s<sup>-1</sup>
- Store the output from each call of the `channel_velocity` function for each value of $n$ in your `channel` DataFrame
    - You should label the new column in your DataFrame using the power-law exponent. For example, you could have a column called `n = 2` for a power-law exponent of 2.
    - **Note**: You should convert your channel velocities from m/s to m/a before storing them in the `channel` DataFrame!
    - The hints for this week's exercise may be helpful for this task

<!-- grade_id: problem_2_ag_p4_velocities, points: 0, grade: false, solution: true -->
```python
# Enter your values below
h = None
y = None
channel = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_2_ag_p4_velocities_test1, points: 1, grade: true, solution: false -->
```python
# These tests should work
print('The shape of channel should be (201, 4). My channel shape is {0}.'.format(channel.shape))
print('The velocity at channel.iloc[51,0] should be 67.1062 m/a. My velocity is {0:.4f}.'.format(channel.iloc[51,0]))

```

## Part 5: Plotting your results (*2 points*)

Next you can create a plot of the velocity calculations you've done.

- Create a plot of the velocity of the ice for all 4 power-law exponents
    - Be sure to label your axes and add a title.
    - Also be sure your line legend makes it clear which profile is which
- Add a figure caption in the Markdown cell beneath the plot, describing the figure as if it were in a scientific journal article

<!-- grade_id: problem_2_mg_p5_plot, points: 1.5, grade: true, solution: true -->
```python
# Add your plot below
ax = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_2_mg_p5_plot_caption, points: 0.5, grade: true, solution: true -->
YOUR ANSWER HERE

# Problem 3 - Comparing our model to the Saskatchewan glacier (*6 points*)

In this problem we will use the `channel_velocity` function from Problem 2 to compare predicted ice velocities to glacier flow velocities observed for the [Saskatchewan glacier near Banff in Alberta, Canada](https://goo.gl/maps/R9S48J4KYPr) (Figure 1). The Saskatchewan glacier is 1400 m wide and part of a large ice field known as the Columbia Icefield.

**Notice**: Closely follow the instructions! For example, you should be sure to use **exactly** the same variable names mentioned in the instructions because your answers will be automatically graded, and the tests that grade your answers rely on following the same formatting or variable naming as in the instructions.

**Your score on this problem will be based on following criteria**:

- Creating a Python function for goodness of fit
- Reading in the data file
- Comparing predictions for the ice flow velocity to observations from the data file
- Plotting your results
- Including comments that explain what most lines in the code do
- Uploading your notebook to your GitHub repository for this week's exercise

![Saskatchewan glacier](https://upload.wikimedia.org/wikipedia/commons/a/a5/Saskatchewan_Glacier.jpg)<br/>
*Figure 1. Saskatchewan glacier.*.

## Part 1: A modified goodness-of-fit function (*1 point*)

The first task for this problem is to create a Python function for calculating a goodness of fit. This will allow us to *quantitatively* compare our ice channel velocity model predictions to data from the Saskatchewan glacier.

For this we will use the same goodness-of-fit equation used back in Exercise 3. As a reminder, the equation for goodness-of-fit is

\begin{equation}
  \Large
  \chi^{2} = \sum \frac{(O_{i} - E_{i})^{2}}{E_{i}}
\end{equation}

where $O_{i}$ is the $i$th observed velocity and $E_{i}$ is the $i$th expected velocity. In this case, the $E_{i}$ values come from our channel velocity function.

- Create a goodness-of-fit function called `chi_squared`

<!-- grade_id: problem_3_ag_p1_function, points: 0, grade: false, solution: true -->
```python
def chi_squared(observed, expected):
    """Returns the chi-squared value for input array data."""
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_3_ag_p1_function_test1, points: 1, grade: true, solution: false -->
```python
# These visible tests with fake data should work
from nose.tools import ok_, assert_equal

# Fake data
obs1 = np.array([1.1, 2.9, 2.6, 3.5, 5.7, 2.8])
exp1 = np.array([1.5, 2.4, 3.6, 1.5, 6.7, 2.6])

obs2 = np.array([1.8, 2.3, 2.2, 3.9, 5.5, 2.4])
exp2 = np.array([1.2, 2.0, 3.9, 1.7, 6.1, 2.9])

# Fake goodness-of-fit values
cs1 = chi_squared(obs1, exp1)
cs2 = chi_squared(obs2, exp2)

# Print results
print("Goodness-of-fit for dataset 1: {0:.4f}.".format(cs1))
print("Goodness-of-fit for dataset 2: {0:.4f}.".format(cs2))

# Check that the chi-squared function works
assert_equal(round(cs1, 4), 3.3199)
assert_equal(round(cs2, 4), 4.0783)
```

## Part 2: Reading the data file (*1 point*)

Distance (m) | Velocity (m/a) | Velocity (m/s)
------------ | -------------- | --------------
-660 | 12 | 3.80E-07
-640 | 22 | 7.00E-07
-570 | 42 | 1.33E-06
-460 | 63 | 2.00E-06
-220 | 74 | 2.35E-06
40 | 76 | 2.41E-06
180 | 74 | 2.35E-06
260 | 72 | 2.28E-06
500 | 51 | 1.62E-06

*Table 1. Velocity measurements across the Saskatchewan glacier.*

The file [`data/sask_glacier_velo.txt`](data/sask_glacier_velo.txt) contains a series of surface velocities measured at various locations across the glacier (Table 1). Your task now is to load in this velocity data.

- Read the data file [`data/sask_glacier_velo.txt`](data/sask_glacier_velo.txt) using pandas into a DataFrame called `sask`
    - Set the column `'Distance (m)'` as the index column

<!-- grade_id: problem_3_ag_p2_read, points: 0, grade: false, solution: true -->
```python
# Load the data file
sask = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_3_ag_p2_read_test1, points: 1, grade: true, solution: false -->
```python
# This test should work
print('The first 5 rows of the data file are\n{0}'.format(sask.head()))

```

## Part 3: Comparing to observations (*2 points*)

The next task is to calculate ice velocities at locations corresponding to the observed velocity measurements and compare the predictions to the data in the data file.

- Loop over the power-law exponent values $n=2$ to $n=5$ and calculate the velocity of ice using the `channel_velocity` function at the positions where the velocity of the Saskatchewan glacier has been observed in the data file
    - Use the default function values other than the power-law exponent and values for `y`
    - You should label the new column in your `sask` DataFrame using the power-law exponent. For example, you could have a column called `n = 2` for a power-law exponent of 2.
    - **Note**: You should convert your channel velocities from m/s to m/a before storing them in the `sask` DataFrame!
- For each set of predicted ice flow velocities, calculate the goodness-of-fit to the observed ice velocities and append this to the `gof_sask` list

<!-- grade_id: problem_3_ag_p3_gof, points: 0, grade: false, solution: true -->
```python
# Calculate the goodness-of-fit
gof_sask = []

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_3_ag_p3_gof_test1, points: 1, grade: true, solution: false -->
```python
# This test should work
print('The first calculated ice velocity should be 12.31 m/a. My first calculated ice velocity is {0:.2f}.'.format(sask.iloc[0,2]))

```

<!-- grade_id: problem_3_ag_p3_gof_test2, points: 1, grade: true, solution: false -->
```python
# This test should work
print('The first goodness of fit should be 3.83. My first goodness of fit is {0:.2f}.'.format(gof_sask[0]))

```

## Part 4: Plotting your results (*2 points*)

The last task in this problem is to plot the predicted ice flow velocities along with the observed values and show the best goodness of fit as text on the plot.

- Plot the measured velocities along with the 4 predicted velocity profiles you have calculated **in meters per year (m/a)**
    - For the predicted velocities, use the values from the `channel` DataFrame from Problem 2
    - For the observed velocities, use the values from the `sask` DataFrame
        - Plot these as black circles
    - Be sure to label your axes and include a title
    - Also add the best goodness-of-fit value as text on the plot along with the corresponding power-law exponent value
- Add a figure caption in the Markdown cell beneath the plot, describing the figure as if it were in a scientific journal article

<!-- grade_id: problem_3_mg_p4_plot, points: 1.5, grade: true, solution: true -->
```python
# Plot your results below
ax = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_3_mg_p4_plot_caption, points: 0.5, grade: true, solution: true -->
YOUR ANSWER HERE

# Problem 4 - Non-Newtonian ice flow down an inclined plane (*7 points*)

The [Athabasca glacier](https://goo.gl/maps/HggYfoKxEUQ2) (Figure 2) is another glacier in the Columbia Icefield, which will be the focus in this problem.

![Athabasca glacier](https://upload.wikimedia.org/wikipedia/commons/4/41/Icefields_parkway.jpg)<br/>
*Figure 2. Athabasca glacier.*

**Notice**: Closely follow the instructions! For example, you should be sure to use **exactly** the same variable names mentioned in the instructions because your answers will be automatically graded, and the tests that grade your answers rely on following the same formatting or variable naming as in the instructions.

**Your score on this problem will be based on following criteria**:

- Creating a Python function for calculating the gravitational force
- Creating a Python function for calculating the ice velocity
- Reading in the data file
- Comparing predictions for the ice flow velocity to observations from the data file
- Plotting your results
- Including comments that explain what most lines in the code do
- Uploading your notebook to your GitHub repository for this week's exercise

## Part 1: Background and theory (*0 points*)

First read through the [background and theory](https://introqg.github.io/site/lessons/L5/exercise-5-theory.html#problem-2) for this problem before proceeding.

- Read through the [background and theory](https://introqg.github.io/site/lessons/L5/exercise-5-theory.html#problem-2) for this problem

## Part 2: A gravitational force function?!? (*1 point*)

In this problem we will be calculating the flow of ice down a slope resulting from a gravitational forces on the glacier. As above, we will consider both Newtonian and non-Newtonian fluids, this time using Equation 19 from the [background and theory for this problem](https://introqg.github.io/site/lessons/L5/exercise-5-theory.html#problem-2), which is also shown below

\begin{equation}
  \large
  u(z) = u_{\mathrm{b}} + \frac{a}{n + 1} \gamma_{x}^{n} \left[ h^{n + 1} - (h - z)^{n + 1} \right]
\end{equation}

where $u_{\mathrm{b}}$ is the basal sliding velocity of the glacier, $a$ is a function of the viscosity and temperature (we will ignore temperature for now), $n$ is the power-law exponent, $\gamma_{x}$ is the downslope component of the gravitational force, $h$ is the thickness of the ice perpendicular to the slope, and $z$ is the array of distances from the base of the glacier. Similar to Problem 3, the value for the gravitational force $\gamma_{x}$ is not known, but we can create a function to estimate its value.

- Create a function called `gravity_force` that calculates the value of $\gamma_{x}^{n}$ for a given value of the power-law exponent `n`.
    - You can solve for $\gamma_{x}^{n}$ by rearranging Equation 19 (see above) to solve for $\gamma_{x}^{n}$ and using the condition that the maximum flow velocity $u = 9.1 \times 10^{-7}$ m/s occurs at $z = h = 209$ m (see Table 2 in Part 4).
    - Note that you should assume the basal sliding velocity $u_{\mathrm{b}}$ is zero for the Newtonian case and equal to the observed sliding velocity $u_{\mathrm{b}} = 1.30 \times 10^{7}$ m/s for the non-Newtonian cases (see Table 2 in Part 4).
    - Again, assume $a = 5 \times 10^{-24}$ Pa<sup>-3</sup> s<sup>-1</sup> and $h = 209$ m.
    - Your function parameters should include
        - The power-law exponent `n`
        - The ice thickness `h`
        - The viscosity/temperature constant `a`
        - The maximum ice flow velocity `u_max`
        - The basal sliding velocity `u_basal`

<!-- grade_id: problem_4_ag_p2_gamma_x, points: 0, grade: false, solution: true -->
```python
def gravity_force(n, h=209.0, a=5.0e-24, u_max=9.1e-7, u_basal=1.3e-7):
    """Calculate gravitational force for ice flowing down a slope.
    
    Keyword arguments:
    n -- viscous flow power-law expoenent
    h -- ice thickness normal to slope (units: m; default: 209.0)
    a -- viscosity/temperature constant (units: Pa^-3 s^-1; default: 5.0e-24)
    u_max -- approximate channel velocity at z = h (units: m/s; default: 9.1e-7)
    u_basal -- basal sliding velocity (units: m/s; default: 1.3e-7)
    """
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_4_ag_p2_gamma_x_test1, points: 0.5, grade: true, solution: false -->
```python
# The visible tests below should work
from nose.tools import assert_equal

# Function test calculations
force1 = gravity_force(n=1, h=209.0, a=5.0e-24, u_max=9.1e-7, u_basal=1.3e-7)

# Print results
print('The gravity force for n=1 should be 8333142556260.158. My gravity force value is {0:.3f}'.format(force1))

# Check the other test calculations
assert_equal(round(force1, 3), 8333142556260.158)
```

<!-- grade_id: problem_4_ag_p2_gamma_x_test2, points: 0.5, grade: true, solution: false -->
```python
# The visible tests below should work
from nose.tools import assert_equal

# Function test calculations
force2 = gravity_force(n=5, h=209.0, a=5.0e-24, u_max=9.1e-7, u_basal=1.3e-7)
force3 = gravity_force(n=7, h=209.0, a=5.0e-24, u_max=9.1e-7, u_basal=1.3e-7)

# Print results
print('The gravity force for n=5 should be 11230.475. My gravity force value is {0:.3f}'.format(force2))

# Check the other test calculations
assert_equal(round(force2, 3), 11230.475)
assert_equal(round(force3, 3), 0.343)
```

## Part 3: An ice flow velocity function, again (*1 point*)

Your next task is to create a function to calculate the velocity profile for ice flowing down a slope (Equation 19), shown in Part 2.

- Create a function called `slope_velocity` that calculates the velocity of ice flowing down a slope using Equation 19 from the [background and theory](https://introqg.github.io/site/lessons/L5/exercise-5-theory.html#problem-2), which was shown in the previous part.
    - Your function should use the `gravity_force` function created in the previous part.
    - As in the previous part you should assume $a = 5 \times 10^{-24}$ Pa<sup>-3</sup> s<sup>-1</sup>, $h = 209$ m, and the maximum flow velocity $u_{\mathrm{max}} = 9.1 \times 10^{-7}$ m/s.
    - Also you should again assume the basal sliding velocity $u_{\mathrm{b}}$ is zero for the Newtonian case and equal to the observed sliding velocity $u_{\mathrm{b}} = 1.30 \times 10^{7}$ m/s for the non-Newtonian cases (see Table 2 in Part 4).
    - Your function parameters should include
        - The power-law exponent `n`
        - The channel thickness array `z`
        - The ice thickness `h`
        - The viscosity/temperature constant `a`
        - The maximum ice flow velocity `u_max`
        - The basal sliding velocity `u_basal`

<!-- grade_id: problem_4_ag_p3_slope_velocity, points: 0, grade: false, solution: true -->
```python
def slope_velocity(n, z, h=209.0, a=5.0e-24, u_max=9.1e-7, u_basal=1.3e-7):
    """Calculate velocity profile of ice flowing downslope.
    
    Keyword arguments:
    n -- viscous flow power-law expoenent
    z -- array of distances from the channel base (units: m)
    h -- ice thickness normal to slope (units: m; default: 209.0)
    a -- viscosity/temperature constant (units: Pa^-3 s^-1; default: 5.0e-24)
    u_max -- Approximate maximum flow velocity at z = h (units: m/s; default: 9.1e-7)
    u_basal -- Basal sliding velocity (units: m/s; default: 1.3e-7)
    """
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_4_ag_p3_slope_velocity_test1, points: 0.5, grade: true, solution: false -->
```python
# The visible tests below should work
from nose.tools import assert_equal

# Function test calculations
z1 = np.array([0.0, 21.0, 74.0, 104.0, 159.0, 209.0])
velocity1 = slope_velocity(n=3, z=z1, h=209.0, a=5.0e-24, u_max=9.1e-7, u_basal=1.3e-7)

# Print results
print('The velocity velocity1[4] for n=3 should be 0.00000091 m/s. My velocity value is {0:.8f}.'.format(velocity1[4]))

# Check the other test calculations
assert_equal(round(velocity1[-1], 8), 9.1e-07)
assert_equal(round(velocity1[3], 8), 8.6e-07)
```

<!-- grade_id: problem_4_ag_p3_slope_velocity_test2, points: 0.5, grade: true, solution: false -->
```python
# The visible tests below should work

# Function test calculations
velocity2 = slope_velocity(n=1, z=z1, h=209.0, a=5.0e-24, u_max=9.1e-7, u_basal=1.3e-7)
velocity3 = slope_velocity(n=7, z=z1, h=209.0, a=5.0e-24, u_max=9.1e-7, u_basal=1.3e-7)

# Print results
print('The velocity velocity2[4] for n=1 should be 0.00000086 m/s. My velocity value is {0:.8f}.'.format(velocity2[4]))

# Check the other test calculations
assert_equal(round(velocity2[-1], 8), 9.1e-07)
assert_equal(round(velocity2[3], 8), 6.8e-07)
assert_equal(round(velocity3[-1], 9), 9.1e-07)
assert_equal(round(velocity3[3], 9), 9.07e-07)
```

## Part 4: Reading the data file (*1 point*)

Depth from surface (m) | Height from base (m) | Velocity (m/a) | Velocity (m/s)
---------------------- | -------------------- | -------------- | --------------
0 | 209 | 28.6 | 9.10E-07
15 | 195 | 28.5 | 9.00E-07
30 | 180 | 28.5 | 9.00E-07
45 | 165 | 28.4 | 9.00E-07
60 | 150 | 28.2 | 8.90E-07
75 | 135 | 28.0 | 8.90E-07
90 | 120 | 27.7 | 8.80E-07
105 | 105 | 27.2 | 8.60E-07
120 | 90 | 26.5 | 8.40E-07
135 | 75 | 25.5 | 8.10E-07
150 | 60 | 24.0 | 7.60E-07
165 | 45 | 21.5 | 6.80E-07
180 | 30 | 17.5 | 5.50E-07
195 | 15 | 10 | 3.20E-07
209 | 0 | 4 | 1.30E-07

*Table 2. Velocities measured from a vertical profile through Athabasca glacier.*

A vertical velocity profile for the Athabasca glacier has been measured and the measurements are in the file [`data/atha_glacier_velo.txt`](data/atha_glacier_velo.txt) (Table 2). Your task now is to load in this velocity data.

- Read the data file [`data/sask_glacier_velo.txt`](data/sask_glacier_velo.txt) using pandas into a DataFrame called `athabasca`
    - **Note**: Do NOT assign an index column, this will cause trouble when plotting the data

<!-- grade_id: problem_4_ag_p4_data, points: 0, grade: false, solution: true -->
```python
# Read the data file below
athabasca = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_4_ag_p4_data_test1, points: 1, grade: true, solution: false -->
```python
# This test should work
print('The first 5 rows of the data file are\n{0}'.format(athabasca.head()))

```

## Part 5: Comparing to observations, again (*2 points*)

The next task is to calculate ice velocities at locations corresponding to the observed velocity measurements and compare the predictions to the data in the data file.

- Loop over the power-law exponent values $n=1$ to $n=5$ and calculate the velocity of ice using the `slope_velocity` function at the positions where the velocity of the Athabasca glacier has been observed in the data file
    - Use the default function values other than the power-law exponent and values for `z`
    - You should label the new column in your `athabasca` DataFrame using the power-law exponent. For example, you could have a column called `n = 2` for a power-law exponent of 2.
    - **Note**: You should convert your ice velocities from m/s to m/a before storing them in the `athabasca` DataFrame!
- For each set of predicted ice flow velocities, calculate the goodness-of-fit to the observed ice velocities and append this to the `gof_athabasca` list

<!-- grade_id: problem_4_ag_p5_gof, points: 0, grade: false, solution: true -->
```python
# Calculate the goodness-of-fit
gof_athabasca = []

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_4_ag_p5_gof_test1, points: 1, grade: true, solution: false -->
```python
# This test should work
print('The first calculated ice velocity should be 28.72 m/a. My first calculated ice velocity is {0:.2f}.'.format(athabasca.iloc[0,4]))

```

<!-- grade_id: problem_4_ag_p5_gof_test2, points: 1, grade: true, solution: false -->
```python
# This test should work
print('The last goodness of fit should be 0.95. My first goodness of fit is {0:.2f}.'.format(gof_athabasca[-1]))

```

## Part 6: Plotting your results (*2 points*)

The last task in this problem is to again plot the predicted ice flow velocities along with the observed values and show the best goodness of fit as text on the plot.

- Plot the measured velocities along with the 5 predicted velocity profiles you have calculated **in meters per year (m/a)** on the x-axis as a function of their **height from the base** on the y-axis
    - For the predicted velocities, use the values from the `athabasca` DataFrame
    - For the observed velocities, also use the values from the `athabasca` DataFrame
        - Plot these as black circles connected by a black line
    - Be sure to label your axes and include a title
    - Also add the best goodness-of-fit value as text on the plot along with the corresponding power-law exponent value
    - **Hint**: You should use the `label` parameter for your plots to provide the text to use for the legend
- Add a figure caption in the Markdown cell beneath the plot, describing the figure as if it were in a scientific journal article

<!-- grade_id: problem_4_mg_p6_plot, points: 1.5, grade: true, solution: true -->
```python
# Plot your results below
ax = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_4_mg_p6_plot_caption, points: 0.5, grade: true, solution: true -->
YOUR ANSWER HERE

# Summary questions for this exercise (*1 point*)

1. Based on your models, what is the most likely power-law exponent for glacial ice?
2. Looking at your plots, how sensitive is the velocity of ice to the power-law exponent?

<!-- grade_id: summary_questions_mg, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE